# Assignment 2 - Network Connectivity

In this assignment you will go through the process of importing and analyzing an internal email communication network between employees of a mid-sized manufacturing company. 
Each node represents an employee and each directed edge between two nodes represents an individual email. The left node represents the sender and the right node represents the recipient. We will also store the timestamp of each email. 

```python
import networkx as nx

#!head assets/email_network.txt
```

### Question 1

Using networkx, load up the directed multigraph from `assets/email_network.txt`. Make sure the node names are strings.

*This function should return a directed multigraph networkx graph.*

```python
import pandas as pd
G_df = pd.read_csv('assets/email_network.txt', sep='\t', 
                   header=0, names=['Sender', 'Recipient', 'time'])
G_df['Sender'] = G_df['Sender'].astype("string")
G_df['Recipient'] = G_df['Recipient'].astype("string")
G_df
```

Output:
```
Sender Recipient        time
0          1         2  1262454010
1          1         3  1262454010
2          1         4  1262454010
3          1         5  1262454010
4          1         6  1262454010
...      ...       ...         ...
82922      3        39  1285868337
82923      3        39  1285868352
82924     19        18  1285877937
82925     85         3  1285884469
82926     85        94  1285884492

[82927 rows x 3 columns]
```

```python
# G = nx.from_pandas_edgelist(G_df, 'Sender', 'Recipient', edge_attr='time')
Graph = nx.MultiDiGraph() 
for i in range(len(G_df)): 
    Graph.add_edge(G_df['Sender'][i], G_df['Recipient'][i], time = G_df['time'][i]) 
```

```python
Graph
```

Output:
```
<networkx.classes.multidigraph.MultiDiGraph at 0x7208a0b7fa90>
```

```python
len(Graph.edges())
```

Output:
```
82927
```

```python
# list(Graph.edges(data=True))
```

<!-- grade_id: cell-60316968695c6f20, points: 0, grade: false, solution: true -->
```python
def answer_one():
    return Graph
```

<!-- grade_id: cell-483a56c89f9db231, points: 1, grade: true, solution: false -->
```python
ans_one = answer_one()
```

### Question 2

How many employees are represented in the network?

How many `sender`->`recipient` pairs of employees are there in the network such that `sender` sent at least one email to `recipient`? Note that even if a `sender` sent multiple messages to a `recipient`, they should only be counted once. You should **not** exclude cases where an employee sent emails to themselves from this [email] count. 

*This function should return a tuple with two integers (#employees, # `sender`->`recipient` pairs).*

```python
# Graph.edges()
```

```python
G_df['Sender'].nunique()
```

Output:
```
154
```

```python
G_df['Recipient'].nunique()
```

Output:
```
140
```

```python
employees = pd.concat([G_df['Sender'], G_df['Recipient']])
employees.nunique()
```

Output:
```
167
```

```python
G_df.groupby(['Sender', 'Recipient']).ngroups
```

Output:
```
5784
```

```python
pairs = []
for tup in Graph.edges():
    if tup not in pairs:
        pairs.append(tup)
len(pairs)
```

Output:
```
5784
```

```python
len(Graph.edges())
```

Output:
```
82927
```

<!-- grade_id: cell-16a78f22207d9f95, points: 0, grade: false, solution: true -->
```python
def answer_two():
    return (167, 5784)
```

<!-- grade_id: cell-5b6391549b076d2b, points: 1, grade: true, solution: false -->
```python
ans_two = answer_two()

```

### Question 3

* Part 1. Assume that information in this company can only be exchanged through email.

    When an employee sends an email to another employee, a communication channel has been created, allowing the sender to provide information to the reciever, but not viceversa. 

    Based on the emails sent in the data, is it possible for information to go from every employee to every other employee?


* Part 2. Now assume that a communication channel established by an email allows information to be exchanged both ways. 

    Based on the emails sent in the data, is it possible for information to go from every employee to every other employee?


*This function should return a tuple of bools (part1, part2).*

```python
nx.is_strongly_connected(Graph)
```

Output:
```
False
```

```python
nx.is_weakly_connected(Graph)
```

Output:
```
True
```

<!-- grade_id: cell-606dd37fb2f55358, points: 0, grade: false, solution: true -->
```python
def answer_three():
    return (False, True)
```

<!-- grade_id: cell-82b3f0bc45e2895f, points: 1, grade: true, solution: false -->
```python
ans_three = answer_three()

```

### Question 4

How many nodes are in the largest weakly connected component of the graph?

*This function should return an int.*

```python
Graph.number_of_nodes()
```

Output:
```
167
```

<!-- grade_id: cell-39196f4c27b35534, points: 0, grade: false, solution: true -->
```python
def answer_four():
    return Graph.number_of_nodes()
```

<!-- grade_id: cell-2b1b7b06ecfa751d, points: 1, grade: true, solution: false -->
```python
ans_four = answer_four()
```

### Question 5

How many nodes are in the largest strongly connected component?

*This function should return an int*

```python
[len(c) for c in sorted(nx.strongly_connected_components(Graph))]
```

Output:
```
[1,
 1,
 1,
 1,
 1,
 1,
 1,
 1,
 1,
 1,
 1,
 1,
 1,
 126,
 1,
 1,
 1,
 1,
 1,
 1,
 1,
 1,
 1,
 1,
 1,
 1,
 1,
 1,
 1,
 1,
 1,
 1,
 1,
 1,
 1,
 1,
 1,
 1,
 1,
 1,
 1,
 1]
```

<!-- grade_id: cell-879d17a9db5c3790, points: 0, grade: false, solution: true -->
```python
def answer_five():
    return 126
```

<!-- grade_id: cell-b0524f7dc1fbdec4, points: 1, grade: true, solution: false -->
```python
ans_five = answer_five()

```

### Question 6

Using the NetworkX functions `strongly_connected_components` and `subgraph`, find the subgraph of nodes in the largest strongly connected component. 
Call this graph G_sc.

*This function should return a networkx MultiDiGraph named G_sc.*

```python
for i in nx.strongly_connected_components(Graph):
    print(i, len(i))
```

Output:
```
{'116'} 1
{'130'} 1
{'133'} 1
{'125'} 1
{'102'} 1
{'126'} 1
{'136'} 1
{'114'} 1
{'127'} 1
{'131'} 1
{'132'} 1
{'135'} 1
{'166'} 1
{'27', '55', '115', '43', '25', '14', '57', '3', '12', '84', '118', '50', '79', '11', '9', '30', '7', '74', '64', '48', '54', '112', '83', '129', '31', '120', '51', '95', '38', '26', '34', '49', '42', '89', '47', '92', '32', '123', '98', '151', '75', '93', '65', '41', '35', '106', '18', '121', '70', '108', '117', '68', '1', '45', '82', '39', '46', '76', '149', '71', '96', '101', '97', '119', '66', '99', '56', '20', '21', '80', '78', '10', '63', '58', '100', '85', '88', '134', '53', '4', '8', '6', '87', '111', '19', '62', '73', '40', '5', '110', '90', '94', '124', '122', '77', '86', '59', '61', '15', '105', '91', '33', '113', '67', '104', '109', '81', '13', '37', '29', '60', '103', '52', '23', '128', '28', '24', '69', '22', '16', '17', '2', '36', '44', '72', '107'} 126
{'137'} 1
{'138'} 1
{'139'} 1
{'140'} 1
{'141'} 1
{'142'} 1
{'143'} 1
{'144'} 1
{'145'} 1
{
```

```python
G_sc = Graph.subgraph({'27', '55', '115', '43', '25', '14', '57', '3', '12', '84', '118', '50', '79', '11', '9', '30', '7', '74', '64', '48', '54', '112', '83', '129', '31', '120', '51', '95', '38', '26', '34', '49', '42', '89', '47', '92', '32', '123', '98', '151', '75', '93', '65', '41', '35', '106', '18', '121', '70', '108', '117', '68', '1', '45', '82', '39', '46', '76', '149', '71', '96', '101', '97', '119', '66', '99', '56', '20', '21', '80', '78', '10', '63', '58', '100', '85', '88', '134', '53', '4', '8', '6', '87', '111', '19', '62', '73', '40', '5', '110', '90', '94', '124', '122', '77', '86', '59', '61', '15', '105', '91', '33', '113', '67', '104', '109', '81', '13', '37', '29', '60', '103', '52', '23', '128', '28', '24', '69', '22', '16', '17', '2', '36', '44', '72', '107'})
```

```python
len(G_sc.nodes())
```

Output:
```
126
```

<!-- grade_id: cell-a81e6f88ef84578c, points: 0, grade: false, solution: true -->
```python
def answer_six():
    return G_sc
```

<!-- grade_id: cell-cf148ef273b3b19c, points: 1, grade: true, solution: false -->
```python
ans_six = answer_six()
assert type(ans_six) == nx.MultiDiGraph , "Your return type should be a MultiDiGraph object"

```

### Question 7

What is the average distance between nodes in G_sc?

*This function should return a float.*

```python
nx.average_shortest_path_length(G_sc)
```

Output:
```
1.6461587301587302
```

<!-- grade_id: cell-1a3075039b3c2c0b, points: 0, grade: false, solution: true -->
```python
def answer_seven():
    return nx.average_shortest_path_length(G_sc)
```

<!-- grade_id: cell-5b374fdd48f37e02, points: 1, grade: true, solution: false -->
```python
ans_seven = answer_seven()

```

### Question 8

What is the largest possible distance between two employees in G_sc?

*This function should return an int.*

```python
ec = nx.eccentricity(G_sc)
max(ec.values())
```

Output:
```
3
```

<!-- grade_id: cell-4fbb8d34d00f675c, points: 0, grade: false, solution: true -->
```python
def answer_eight():
    return 3
```

<!-- grade_id: cell-c5714787854ef644, points: 1, grade: true, solution: false -->
```python
ans_eight = answer_eight()

```

### Question 9

What is the set of nodes in G_sc with eccentricity equal to the diameter?

*This function should return a set of the node(s).*

```python
for key, value in ec.items():
    if value==3:
        print(key)
```

Output:
```
97
129
134
```

<!-- grade_id: cell-bcfab850744277ed, points: 0, grade: false, solution: true -->
```python
def answer_nine():
    return {'97', '129', '134'}
```

<!-- grade_id: cell-77c9ca0b94df3d6f, points: 1, grade: true, solution: false -->
```python
ans_nine = answer_nine()
assert type(ans_nine) == set, "Student answer must return a set"
```

### Question 10

What is the set of node(s) in G_sc with eccentricity equal to the radius?

*This function should return a set of the node(s).*

```python
for key, value in ec.items():
    if value==min(ec.values()):
        print(key)
```

Output:
```
38
```

<!-- grade_id: cell-fd3d43ea37da03f7, points: 0, grade: false, solution: true -->
```python
def answer_ten():
    return {'38'}
```

<!-- grade_id: cell-bfd2ee304bc25264, points: 1, grade: true, solution: false -->
```python
ans_ten = answer_ten()
assert type(ans_ten) == set, "Student answer must return a set"
```

### Question 11

Which node in G_sc has the most shortest paths to other nodes whose distance equal the diameter of G_sc?


For the node with the most such shortest paths, how many of these paths are there?


*This function should return a tuple (name of node, number of paths).*

```python
# nx.shortest_path_length(G_sc, source='1')
```

```python
num_path = {}
for node in G_sc.nodes():
    count = 0
    path = nx.shortest_path_length(G_sc, source=node)
    for i in path:
        if path[i]==3:
            count = count + 1
    num_path[node] = count
```

```python
print(max(num_path.values()))
for key in num_path:
    if num_path[key] == max(num_path.values()):
        print(key)
```

Output:
```
63
97
```

<!-- grade_id: cell-c40858baf3d92139, points: 0, grade: false, solution: true -->
```python
def answer_eleven():
    return ('97', 63)
```

<!-- grade_id: cell-f79b06650f61cf37, points: 1, grade: true, solution: false -->
```python
ans_eleven = answer_eleven()
assert type(ans_eleven) == tuple, "Student answer must be a tuple"
```

### Question 12

Suppose you want to prevent communication flow from the node that you found in question 11 to node 10. What is the smallest number of nodes you would need to remove from the graph (you're not allowed to remove the node from the previous question or 10)? 

*This function should return an integer.*

```python
nx.minimum_node_cut(G_sc, '97', '10')
```

Output:
```
{'14'}
```

```python
node_cut = nx.minimum_node_cut(G_sc, '97', '10')
len(node_cut)
```

Output:
```
1
```

<!-- grade_id: cell-a4606db8971dcbba, points: 0, grade: false, solution: true -->
```python
def answer_twelve():
    return 1
```

<!-- grade_id: cell-509cfa9f4136124d, points: 1, grade: true, solution: false -->
```python
ans_twelve = answer_twelve()

```

### Question 13

Convert the graph G_sc into an undirected graph by removing the direction of the edges of G_sc. Call the new graph G_un. 


*This function should return a networkx Graph.*

```python
G_mu = G_sc.to_undirected()
```

```python
G_mu
```

Output:
```
<networkx.classes.multigraph.MultiGraph at 0x72089e8abb80>
```

```python
# G_mu.edges()
```

```python
G_un = nx.Graph()
for u,v in G_mu.edges():
    if G_mu.has_edge(u,v):
        G_un.add_edge(u, v)
```

```python
G_un
```

Output:
```
<networkx.classes.graph.Graph at 0x72089e86a130>
```

<!-- grade_id: cell-612b40425f0851e8, points: 0, grade: false, solution: true -->
```python
def answer_thirteen():
    return G_un
```

<!-- grade_id: cell-d1c0627a327cd774, points: 1, grade: true, solution: false -->
```python
ans_thirteen = answer_thirteen()
assert type(ans_thirteen) == nx.Graph , "Your return type should be a Graph object"

```

### Question 14

What is the transitivity and average clustering coefficient of graph G_un?

*This function should return a tuple (transitivity, avg clustering).*     
*Note: DO NOT round up your answer.*

```python
print(nx.transitivity(G_un))
```

Output:
```
0.570111160700385
```

```python
print(nx.average_clustering(G_un))
```

Output:
```
0.6975272437231419
```

<!-- grade_id: cell-63e48b1eed1a1fb8, points: 0, grade: false, solution: true -->
```python
def answer_fourteen():
    return (nx.transitivity(G_un), nx.average_clustering(G_un))
```

<!-- grade_id: cell-41dda91202f58e7e, points: 1, grade: true, solution: false -->
```python
ans_fourteen = answer_fourteen()
assert type(ans_fourteen) == tuple, "Student answer must be a tuple"

```

```python

```